---
title: 'Les Objets'
description: 'Ce chapitre explique le concept fondamental de l''objet dans R'
---

## Le concept d'objet

```yaml
type: VideoExercise
key: db28a1b17b
xp: 50
```

`@projector_key`
09d3f00d0ec72acb5e1bcb1b64ded46a

---

## Exercice sur le nom de l'objet

```yaml
type: PureMultipleChoiceExercise
key: b46c2b3167
xp: 50
```

Parmi les mots suivants lesquels sont des mots possibles pour un objet ?

`@hint`
Rappelez vous : une lettre, uniquement des lettres, des chiffres, et les caracteres "_", "-" et "."

`@possible_answers`
1. [monObjet]
2. mon objet
3. [mon_objet]
4. 1objet
5. l'objet
6. [l.objet]

`@feedback`
1. Oui c'est correct
2. NON il ne faut pas d'espaces
3. Oui le trait du bas est autorisé
4. le nom d'objet ne doit pas commencer par un chiffre
5. l'apostrophe ne peut pas faire parti du nom de l'objet
6. Oui le point est autorisé

---

## Insert exercise title here

```yaml
type: NormalExercise
key: fb10cc4405
xp: 100
```

Exécuter la première ligne :
```{r, eval=TRUE}
monObjetNum <- 122
```

rappel : cliquez sur la ligne et rajouter la valuer puis tapez entrée en maintenant cmd enfoncé
ou retapez la ligne dans la console (plus long et risque d'erreur)


`@instructions`
# Code à modifier ci-dessous :
monObjetNum <-

`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}
monObjetNum <-
```

`@solution`
```{r}
monObjetNum <- 122
```

`@sct`
```{r}

```
