---
title: 'Chapter 1'
description: 'Prise en main de R'
---

## Example coding exercise

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

Executez les commandes qui apparaissent dans la partie script

`@instructions`
Pour cela, au choix :

séléctionner la première ligne et faites commande-entrez
retaper la ligne dans le quadrant de la console
cliquer sur l'icone “Run”
cliquer sur l'icone “Source” (execute toutes les lignes)

`@hint`


`@pre_exercise_code`
```{r}
"Bonjour"
32 * 44
2.3 + 4.201
```

`@sample_code`
```{r}
"Bonjour"
32 * 44
2.3 + 4.201
```

`@solution`
```{r}

```

`@sct`
```{r}

```
